<problem display_name="Derivative (uses mitxgraders zip)" markdown="&lt;problem&gt;&#10;&lt;script type=&quot;text/python&quot;&gt;&#10;from mitxgraders import *&#10;grader = FormulaGrader(&#10;    variables=[&quot;x&quot;],&#10;    # allows students to use generic functions f and f' in their input&#10;    user_functions={&quot;f&quot;: RandomFunction(), &quot;f'&quot;: RandomFunction()}&#10;)&#10;&lt;/script&gt;&#10;&#10;  &lt;p&gt;Enter the derivative of \(g(x) = e^{f(x^2)} \).&lt;/p&gt;&#10;  &lt;!-- answer is provided to the grader when using single inputs --&gt;&#10;  &lt;customresponse cfn=&quot;grader&quot; answer=&quot;e^(f(x)) * f'(x^2) * 2*x&quot;&gt;&#10;    &lt;textline math=&quot;true&quot; /&gt;&#10;  &lt;/customresponse&gt;&#10;&#10;&lt;/problem&gt;">
<problem>
<script type="text/python">
from mitxgraders import *
grader = FormulaGrader(
    variables=["x"],
    # allows students to use generic functions f and f' in their input
    user_functions={"f": RandomFunction(), "f'": RandomFunction()}
)
</script>

  <p>Enter the derivative of \(g(x) = e^{f(x^2)} \).</p>
  <!-- answer is provided to the grader when using single inputs -->
  <customresponse cfn="grader" answer="e^(f(x)) * f'(x^2) * 2*x">
    <textline math="true"/>
  </customresponse>

</problem>
</problem>
